# Lab 3

1. 提交作業之前，建議可以先點選上方工具列的**Kernel**，再選擇**Restart & Run All**，檢查一下是否程式跑起來都沒有問題，最後記得儲存。
2. 請先填上下方的姓名(name)及學號(stduent_id)再開始作答，例如：
```python
name = "我的名字"
student_id= "B06201000"
```
3. 演算法的實作可以參考[lab-3](https://yuanyuyuan.github.io/itcm/lab-3.html), 有任何問題歡迎找助教詢問。
4. **Deadline: 10/30(Wed.)**

```python
name = ""
student_id = ""
```

---

# Exercise 1

## Let $g(x)=\ln(4+x-x^2)$ and $\alpha$ is a fixed point of $g(x)$ i.e. $\alpha = g(\alpha).$ 

- ### Part A. Implement your fixed-point algorithm and solve it with initial guess $x_0=2$ within tolerance $10^{-10}$, and answer the questions of error behavior analysis below.
- ### Part B. Redo Part A. by applying Aitken's acceleration.

---

### Import libraries

```python
import numpy as np
import matplotlib.pyplot as plt
```

### Implement the target function $g(x)=\ln(4+x-x^2)$

<!-- grade_id: cell-07da59b434940496, points: 0, grade: false, solution: true -->
```python
def g(x):
    # ===== 請實做程式 =====
    
    # ====================
```

<!-- grade_id: cell-c0f08330aec65e17, points: 1, grade: true, solution: false -->
```python
assert round(g(0), 4) == 1.3863
```

### Run built-in [fixed-point method](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.fixed_point.html#rf001026ae057-1) with Python SciPy, and use this accurate value as the fixed point $\alpha$

```python
from scipy import optimize
alpha = optimize.fixed_point(g, x0=2, xtol=1e-12)
print('The fixed point is', alpha)
```

### Visualization

```python
x_range = np.arange(1, 2.25, 0.01)
plt.figure(figsize=(16, 9))
plt.title(r'Solve $y=g(x)$')
plt.plot(x_range, g(x_range), label=r'$y=f(x)$')
plt.plot(x_range, x_range, label=r'$y=x$')
plt.plot(alpha, g(alpha), 'bo', label='fixed point')
plt.plot(2.0, g(2.0), 'r^', label='starting point')
plt.axvline(x=alpha, linestyle='--', label=r'$y=\alpha$')
plt.gca().legend()
plt.grid()
plt.show()
```

---

## Part A.

### 1. Find the fixed point of $g(x)$ using your fixed-point iteration to within tolerance $10^{-10}$ with initial guess $x_0=2$.

#### 1-1. Implement the fixed point method

<!-- grade_id: cell-a6170a4c3abf32bb, points: 0, grade: false, solution: true -->
```python
def fixed_point(
    func,
    x_0,
    tolerance=1e-7,
    max_iterations=5,
):
    '''Find the fixed point of the given function func

    Parameters
    ----------
    func : function
        The target function.
    x_0 : float
        Initial guess point for a solution func(x)=x.
    tolerance: float
        One of the termination conditions. Error tolerance.
    max_iterations : (positive) integer
        One of the termination conditions. The amount of iterations allowed.

    Returns
    -------
    solution : float
        Approximation of the root.
    history: dict
        Return history of the solving process
        history: {'x_n': list}
    '''

    # ===== 請實做程式 =====
    
    # ====================
```

#### 1-2. Find the root

<!-- grade_id: cell-310762e452216efc, points: 0, grade: false, solution: true -->
```python
solution, history = fixed_point(
    # ===== 請實做程式 =====
    
    # ====================
)
```

<!-- grade_id: cell-2d72f68109ee500c, points: 1, grade: true, solution: false -->
```python
print('My estimation is', solution)
```

### 2. Estimate graphically the asymptotic error constant C 
### $$\lim_{n\to\infty} \frac{|x_{n+1} - \alpha|}{|x_{n} - \alpha|} = C$$

<!-- grade_id: cell-6a067de2a49274b5, points: 0, grade: false, solution: true -->
```python
'''
Hint:
    1. Prepare the sequences: x_n(from the history of algorithm)
    2. Compute the error of sequence: e_n
    3. Compute the sequence: e_{n+1}/e_{n}
    4. Plot the curve
    5. Fill in the name of x,y axes
    6. Show the plot
''' 
# ===== 請實做程式 =====

# ====================
```

---

## Part B.

### 1. Accelerate the convergence of the sequence $\{x_n\}$ obtained in _Part A._ using Aitken's  $\Delta^2$ method, yielding sequence $\{\hat{x}_n\}$. 

#### 1-1. Introduce Aitken's acceleration into the original method.

<!-- grade_id: cell-d90669ed0519db80, points: 0, grade: false, solution: true -->
```python
def aitken(
    func,
    x_0,
    tolerance=1e-7,
    max_iterations=5,
):
    '''Approximate solution of f(x)=0 on interval [a,b] by the secant method.

    Parameters
    ----------
    func : function
        The target function.
    x_0 : float
        Initial guess point for a solution f(x)=x.
    tolerance: float
        One of the termination conditions. Error tolerance.
    max_iterations : (positive) integer
        One of the termination conditions. The amount of iterations allowed.

    Returns
    -------
    solution : float
        Approximation of the root.
    history: dict
        Return history of the solving process
        history: {'x_n': list}
    '''

    # ===== 請實做程式 =====
    
    # ====================
```

#### 1-2. Find the root

<!-- grade_id: cell-0767f880d4679e22, points: 0, grade: false, solution: true -->
```python
solution, history = aitken(
    # ===== 請實做程式 =====
    
    # ====================
)
```

<!-- grade_id: cell-5c862e35ba0aa7d9, points: 1, grade: true, solution: false -->
```python
print('My estimation is', solution)
```

### 2. Plot the error curves of each algorithm w.r.t iterations $n$ in log scale to compare the convergence rates. You may see a figure like the one in our lecture.

<figure>
  <img src="aitken.png" style="width:60%">
    <figcaption>Ref. Page15 of <a href="https://ceiba.ntu.edu.tw/course/7a770d/content/cmath2019_note1_aitken.pdf">cmath2019_note1_aitken.pdf</a></figcaption>
</figure>

<!-- grade_id: cell-6f495072de649c1e, points: 0, grade: false, solution: true -->
```python
'''
Hint:
    1. Prepare the sequences: x_n, x_n_hat(from the history of each algorithm)
    2. Compute the error of sequences: e_n, e_n_hat
    3. Plot the curves of e_n, e_n_hat respectively
    4. Change scale into log
    5. Fill in the name of x,y axes
    6. Enable legend(show curve names)
    7. Show the plot
''' 
# ===== 請實做程式 =====

# ====================
```